## Problem Statment

You have been provided a dataset that consists of echo-location clicks of two types of whales, namely, Gervais and Cuviers. In this assignment, your task is to classify the different types whales using Gradient Boosting with the help of the XGBoost library. You are expected to fill in functions that would complete this task. We use XGBoost here instead of GradientBoostedTrees in Spark because XGBoost running on a single machine is much faster than Spark running on 10 machines.

The data files were preprocessed on PySpark (10 nodes) cluster. The code for the same can be found in Data_Processing_Whales.ipynb. The preprocessed data is a numpy array with `4175` rows (for the 10mb file) with following columns (zero-indexed):
* Col 0-9: projections on first 10 eigen vectors
* Col 10: rmse
* Col 11: peak2peak
* Col 12: label (`0 if row.species==u'Gervais' else 1`)

You can also refer to XGBoost_Whales.ipynb under for more details on the XGBoost Analysis before you attempt this assignment.

Both Data_Processing_Whales.ipynb and XGBoost_Whales.ipynb can be found under XGBoost directory that was uploaded in edX as a part of "Notebooks for weeks 7 & 8". 

## XGBoost - Theory

A brief overview of gradient boosting in XGBoost can be found here:

* http://xgboost.readthedocs.io/en/latest/model.html
* http://xgboost.readthedocs.io/en/latest/python/python_intro.html



Use the XGBoost API for training and predicting scores: 

* http://xgboost.readthedocs.io/en/latest/python/python_api.html

#### Main API

* `xgboost.train` is the learning API that trains the Gradient Boosting Model,
   * The main parameters are:
      * **plst** – XGBoost parameter list
      * **dtrain** – Data to be trained
      * **num_round** – Number of iterations of boosting. (default: 100)
      * **evallist** – List of items to be evaluated during training
      * **verbose_eval** - This can be used to control how much information the train function prints. You might want to set to **False** to avoid printing logs.
* `bst.predict` is the API that makes score predictions
   * The main parameters are:
      * **dtest** – Test Data
      * **dtrain** – Data to be trained
      * **ntree_limit** – Limit number of trees in the prediction (Use: ntree_limit=bst.best_ntree_limit)
      * **output_margin** - Whether to output the raw untransformed margin value (Use: output_margin=True)

## Notebook Setup

### Importing Required Libraries

```python
%matplotlib inline
import numpy as np
import xgboost as xgb
from sklearn.model_selection import train_test_split
import matplotlib.pyplot as plt
import pickle
import random
```

### Loading Data

```python
with open('Data/X_train.pkl', 'rb') as f:
    X_train = pickle.load(f)

with open('Data/X_test.pkl', 'rb') as f:
    X_test = pickle.load(f)

with open('Data/y_train.pkl', 'rb') as f:
    y_train = pickle.load(f)

with open('Data/y_test.pkl', 'rb') as f:
    y_test = pickle.load(f)
```

### Setting Parameters for XG Boost
* Maximum Depth of the Tree = 3 _(maximum depth of each decision trees)_
* Step size shrinkage used in update to prevents overfitting = 0.3 _(how to weigh trees in subsequent iterations)_
* Evaluation Criterion= Maximize Loglikelihood according to the logistic regression _(logitboost)_
* Maximum Number of Iterations = 1000 _(total number trees for boosting)_
* Early Stop if score on Validation does not improve for 5 iterations

[Full description of options](https://xgboost.readthedocs.io/en/latest//parameter.html)

```python
#You can change this cell if you wish to, but you aren't expected to
def xgboost_plst():
    param = {}
    param['max_depth']= 3   # depth of tree
    param['eta'] = 0.3      # shrinkage parameter
    param['silent'] = 1     # not silent
    param['objective'] = 'binary:logistic'
    param['nthread'] = 7 # Number of threads used
    param['eval_metric'] = 'logloss'

    plst = param.items()
    return plst
```

## Exercises

### Computing the score ranges

The function <font color="blue">calc_stats</font> takes the xgboost margin scores as input and returns two numpy arrays min_scr, max_scr which are calculated as follows:

1. **min_scr**: mean - (3 $\times$ std)
2. **max_scr**: mean + (3 $\times$ std)

Here the input margin scores, represents the processed XGBoost margin scores obtained from the <font color="blue">bootstrap_pred</font> function. Each row represents the various scores for a specific example in an iteration and your <font color="blue">calc_stats</font> function is supposed to compute the **min_scr** and **max_scr** as defined for each example. So in the example below, we take a scenario where we have 3 examples which have 4 values each (From 4 bootstrap iterations).


**<font color="magenta" size=2>Example Input</font>**
``` python
[[-0.22 -0.19 -0.17 -0.13][-0.1 -0.05 0.02 0.10][0.03 0.11 0.12 0.15]]
```

Output: min_scr (numpy array), max_scr (numpy array)

**<font color="blue" size=2>Example Output</font>**
``` python
(array([-0.28 -0.23 -0.03]),
 array([-0.08  0.22  0.24]))
```

**Note**: Ensure you round the values in the numpy arrays to two decimal places

<!-- grade_id: cell-61aa63fe2f976d34, points: 0, grade: false, solution: true -->
```python
def calc_stats(margin_scores):
    ###
    ### YOUR CODE HERE
    ###
    mean = np.mean(margin_scores, axis = 1)
    std = np.std(margin_scores, axis = 1)
    return np.around(mean - 3 * std, 2), np.around(mean + 3 * std, 2)
```

<!-- grade_id: calc_stats_rt, points: 4, grade: true, solution: false -->
```python
margin_score = np.array([[-0.22, -0.19, -0.17, -0.13], [-0.1, -0.05, 0.02, 0.10], [0.03, 0.11, 0.12, 0.15]])
min_score, max_score = calc_stats(margin_score)
assert type(min_score) == np.ndarray, 'Incorrect Return type'
assert type(max_score) == np.ndarray, 'Incorrect Return type'
```

<!-- grade_id: calc_stats_v1, points: 4, grade: true, solution: false -->
```python
assert (min_score == np.array([-0.28, -0.23, -0.03])).all(), "Incorrect return value"
```

<!-- grade_id: calc_stats_v2, points: 4, grade: true, solution: false -->
```python
assert (max_score == np.array([-0.08,  0.22,  0.24])).all(), "Incorrect return value"
```

<!-- grade_id: calc_stats_h1, points: 4, grade: true, solution: false -->
```python
# Hidden Tests Here
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: calc_stats_h2, points: 4, grade: true, solution: false -->
```python
# Hidden Tests Here
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

### Calculating predictions

Based on the ranges for each of the examples, i.e, (min_scr, max_scr), we can predict whether it's a Gervais or a Cuvier. Since all our scores will be between -1 and +1, we use 0 as the margin line. All examples which are on the left side of the margin, can be classified as Cuvier's and all which are on the right side can be classified as Gervais. However, since we take margin scores from a set of bootstraps for each example, we use the minimum and maximum score arrays to predict to determine the classification.

The function <font color="blue">predict</font> takes the minimum score array and maximum score array and returns predictions as follows:

1. If respective minimum score and maximum score values are less than 0, predict -1 (**Cuvier's**)
2. If respective minimum score value is less than 0 and maximum score value is greater than 0, predict 0 (**Unsure**)
3. If respective minimum score and maximum score values are greater than 0, predict 1 (**Gervais**)

**<font color="magenta" size=2>Example Input</font>**
``` python
min_scr (numpy array) = [-0.78 -0.68 -0.6 -0.53 -0.47 -0.42 -0.32 -0.21 -0.07 0.22]

max_scr (numpy array) = [-0.49 -0.39 -0.33 -0.25 -0.2 -0.11 -0.04 0.1 0.3 0.51]
```
Output: pred (numpy array of predictions)

**<font color="blue" size=2>Example Output</font>**
``` python
[-1 -1 -1 -1 -1 -1 -1  0  0  1]
```

<!-- grade_id: cell-ca97c577ab4f263b, points: 0, grade: false, solution: true -->
```python
def predict(min_scr, max_scr):
    ###
    ### YOUR CODE HERE
    ###
    min_predict = np.where(min_scr < 0, -1, 0)
    max_predict = np.where(max_scr < 0, 0, 1)
    return min_predict + max_predict
```

```python
max_s = np.array([-0.49, -0.39, -0.33, -0.25, -0.2, -0.11, -0.04, 0.1, 0.3, 0.51])
min_s = np.array([-0.78, -0.68, -0.6, -0.53, -0.47, -0.42, -0.32, -0.21, -0.07, 0.22])
pred = predict(min_s, max_s)
true_pred = np.array([-1, -1, -1, -1, -1, -1, -1, 0, 0, 1])

```

<!-- grade_id: predictions_rt, points: 4, grade: true, solution: false -->
```python
assert type(pred) == np.ndarray, 'Incorrect return type'
```

<!-- grade_id: predictions_v1, points: 4, grade: true, solution: false -->
```python
assert (pred == true_pred).all(), 'Incorrect return value'
```

### Calculating scores

The function <font color="blue">bootstrap_pred</font> takes as input:

1. **Training set**
1. **Test set**
1. **n_bootstrap** Number of bootstrap samples that run XGBoost and trains one part of the sample set.
1. **minR, maxR** two numbers such that $0 < minR < maxR < 1$ that define the fractions of the `n_bootstrap` scores that define the range.
1. **bootstrap_size** - Number of bootstrap samples on which you will run XGBoost.
1. **num_round** - Number of iterations for running xgboost
1. **plst** - Parameter List

The output should be a confidence interval for each example in the test set. Together with a prediction that is `Gervais / Cuviers / Unsure`. The prediction `unsure` is to be output if the confidence interval contains the point 0.
After generating the confidence intervals, the function <font color="blue">predict</font> can be used to make predictions.

**Procedure**

Repeat the given procedure for n_bootstrap number of iterations:

For **n_bootstrap** iterations:
* Sample **boostrap_size** indices from the training set **with replacmennt**
* Create train and test data matrices (dtrain, dtest) using xgb.DMatrix(X_sample, label=y_sample)
* Initialise the evallist parameter [(dtrain, 'train'), (dtest, 'eval')]
* Train the model using the XGBoost train API and make score predictions using bst.predict object returned by XGB train API. Ensure you set **output_margin=True** to get raw untransformed output scores.
* Normalize them by dividing them with the normalizing factor as max(scores) - min(scores) and round these values to a precision of two decimal places

Then: 
* For each individual example, remove scores below the minRth percentile and greater than maxRth percentile (sort for each example if necessary)
* Call the calc_stats function to compute min_scr and max_scr with the filtered margin scores as parameter
* Return the min_scr and max_scr computed by the **calc_stat** function using the margin scores.
      
**Note**: You can experiment by changing **n_bootstraps**, but it takes about 200 iterations to get consistent values.

<!-- grade_id: cell-1140639583a712cb, points: 0, grade: false, solution: true -->
```python
def bootstrap_pred(X_train, X_test, y_train, y_test, n_bootstrap, minR, maxR, bootstrap_size, \
                   num_round=100, plst=xgboost_plst()):
    ###
    ### YOUR CODE HERE
    ###
    all_preds = []
    for i in range(n_bootstrap):
        indices = np.random.choice(X_train.shape[0], bootstrap_size)
        X_sample = X_train[indices, :]
        y_sample = y_train[indices]
        dtrain = xgb.DMatrix(X_sample, label=y_sample)
        dtest = xgb.DMatrix(X_test, label=y_test)
        evallist = [(dtrain, 'train'), (dtest, 'eval')]
        bst = xgb.train(plst, dtrain, num_round, evallist)
        preds = bst.predict(dtest, output_margin=True)
        preds = np.around(preds/(max(preds) - min(preds)), 2)
        all_preds.append(preds)
    all_preds = np.array(all_preds).T
    lower_bound = int(n_bootstrap*minR)
    higher_bound = int(n_bootstrap*maxR)
    all_preds = all_preds[:,lower_bound: higher_bound]
    min_scr, max_scr = calc_stats(all_preds)
    return min_scr,max_scr
```

```python
def process(X_train, X_test, y_train, y_test, n_bootstrap=100):
    min_scr, max_scr = bootstrap_pred(X_train, X_test, y_train, y_test, n_bootstrap=n_bootstrap, \
                                            minR=0.1, maxR=0.9, bootstrap_size=len(X_train))
    pred = predict(min_scr, max_scr)
    return min_scr, max_scr, pred
```

#### Tests

How we test the function:
1. We have computed the average mid-point of the range of values and verify that this midpoint is present in the range computed by your function
2. We check that the length of the interval(max_scr-min_scr) is not more than twice the average length of the interval

```python
sample_indices = np.load('Data/vis_indices.npy')
X_test_samp = X_test[sample_indices]
y_test_samp = np.array(y_test[sample_indices], dtype=int)
midpt = np.load('Data/vis_midpt.npy')
avg_length = np.load('Data/vis_avg_length.npy')
min_scr, max_scr, pred = process(X_train, X_test_samp, y_train, y_test_samp)
length = max_scr - min_scr
```

Output:
```
[0]	train-logloss:0.556681	eval-logloss:0.627987
[1]	train-logloss:0.479583	eval-logloss:0.584117
[2]	train-logloss:0.42965	eval-logloss:0.593812
[3]	train-logloss:0.394363	eval-logloss:0.556308
[4]	train-logloss:0.371759	eval-logloss:0.573538
[5]	train-logloss:0.352995	eval-logloss:0.545587
[6]	train-logloss:0.338188	eval-logloss:0.564343
[7]	train-logloss:0.323418	eval-logloss:0.56144
[8]	train-logloss:0.31383	eval-logloss:0.57312
[9]	train-logloss:0.304848	eval-logloss:0.595178
[10]	train-logloss:0.296081	eval-logloss:0.615124
[11]	train-logloss:0.289182	eval-logloss:0.632748
[12]	train-logloss:0.284802	eval-logloss:0.625966
[13]	train-logloss:0.277771	eval-logloss:0.673662
[14]	train-logloss:0.272437	eval-logloss:0.696293
[15]	train-logloss:0.267951	eval-logloss:0.710683
[16]	train-logloss:0.262874	eval-logloss:0.68156
[17]	train-logloss:0.258881	eval-logloss:0.704092
[18]	train-logloss:0.254952	eval-logloss:0.693478
[19]	train-logloss:0.249379	eval-logloss:0.682459
[20]	train-logl
```

<!-- grade_id: process_v1, points: 4, grade: true, solution: false -->
```python
assert sum(min_scr <= midpt) >= (0.7 * len(sample_indices)), "Incorrect range (mean - 3*std) to (mean + 3*std)"
```

<!-- grade_id: process_v2, points: 4, grade: true, solution: false -->
```python
assert sum(max_scr >= midpt) >= (0.7 * len(sample_indices)), "Incorrect range (mean - 3*std) to (mean + 3*std)"
```

<!-- grade_id: process_v3, points: 4, grade: true, solution: false -->
```python
assert sum(length < 2*avg_length) >= (0.7 * len(sample_indices)), "Incorrect length of range (mean - 3*std) to (mean + 3*std)"
```

<!-- grade_id: process_h1, points: 8, grade: true, solution: false -->
```python
# Hidden Tests Here
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: process_h2, points: 8, grade: true, solution: false -->
```python
# Hidden Tests Here
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: process_h3, points: 8, grade: true, solution: false -->
```python
# Hidden Tests Here
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: process_h4, points: 6, grade: true, solution: false -->
```python
# Hidden Tests Here
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: process_h5, points: 8, grade: true, solution: false -->
```python
# Hidden Tests here
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: process_h6, points: 8, grade: true, solution: false -->
```python
# Hidden Tests here
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: process_h7, points: 8, grade: true, solution: false -->
```python
# Hidden Tests here
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: process_h8, points: 6, grade: true, solution: false -->
```python
# Hidden Tests here
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```